<div class="alert alert-block alert-info">
<b>Number of points for this notebook:</b> 0.5
<br>
<b>Deadline:</b> March 2, 2020 (Monday). 23:00
</div>

# Exercise 1.3. Multilayer perceptron with multiple hidden layers

The goal of this exercise is to get familiar with the basics of PyTorch and train a multilayer perceptron (MLP) model.

```python
skip_training = False  # Set this flag to True before validation and submission
```

<!-- grade_id: evaluation_settings, points: 0, grade: true, solution: false -->
```python
# During evaluation, this cell sets skip_training to True
# skip_training = True
```

```python
import os
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
%matplotlib inline

import torch
import torch.nn as nn
import torch.nn.functional as F

import tools
import data
```

```python
if skip_training:
    # The models are always evaluated on CPU
    device = torch.device("cpu")
```

# Multilayer perceptron (MLP) network with multiple hidden layers

Your task is to define a multilayer perceptron with an arbitrary number of layers and arbitrary number of neurons in each layer, so that an MLP can be created as follows:
```python
mlp = MLP([11, 150, 100, 50, 2], activation_fn=F.tanh)
```
In the example above, we created a network with 11 inputs and 2 outputs and three hidden layers: 150 units in the first hidden layer, 100 units in the second one and 50 units in the third one.

Note: The same activation function should be applied to all the layers except for the last one. This way the MLP can be used either for regression or classification.

Hint:
* You can use functions [`torch.nn.Module.add_module`](https://pytorch.org/docs/master/nn.html#torch.nn.Module.add_module) or class [`torch.nn.ModuleList`](https://pytorch.org/docs/stable/nn.html#torch.nn.ModuleList) to register parameters of the model.
* **We recommend you to create an MLP with your implementation and train (please do not have the training loop in the submitted version**).

<!-- grade_id: MLP, points: 0, grade: false, solution: true -->
```python
class MLP(nn.Module):
    def __init__(self, sizes, activation_fn=torch.tanh):
        """Multilayer perceptron with an arbitrary number of layers.
        
        Args:
          sizes (list): Number of units in each layer including the input and the output layer:
                         [n_inputs, n_units_in_hidden_layer1, ..., n_units_in_hidden_layerN, n_outputs]
          activation_fn (callable): An element-wise function used in every layer except in the last one.
        """
        super(MLP, self).__init__()
        self.linears = nn.ModuleList([torch.nn.Linear(sizes[i], sizes[i+1]) for i in range(len(sizes)-1)])
        self.activation_fn = activation_fn
    def forward(self, x):
        for i, l in enumerate(self.linears):
            if i==len(self.linears):
                x=l(x)
                continue
            x = self.activation_fn(l(x))
        return x
```

<!-- grade_id: cell-1dcdd8e8bdecf07e, points: 0.5, grade: true, solution: false -->
```python
# Let us create the network and make sure it can process a random input of the right shape
def test_MLP_shapes():
    n_inputs = 11
    n_samples = 10
    net = MLP([n_inputs, 100, 50, 2])
    y = net(torch.randn(n_samples, n_inputs))
    assert y.shape == torch.Size([n_samples, 2]), f"Bad y.shape={y.shape}"
    print('Success')

test_MLP_shapes()
```

Output:
```
Success
```

```python

```